Gut microbiota-derived gamma-aminobutyric acid from metformin treatment reduces hepatic ischemia/reperfusion injury through inhibiting ferroptosis

Author Response

The following is the authors’ response to the previous reviews.

Reviewer #1 (Recommendations For The Authors):

Many of my specific issues have been addressed in the revision. However, the data shown in Reviewer Fig. 1 and 2 is not sufficiently described to assess it's reliability and these new data do not appear to have been integrated into the paper. A response that more clearly states how the manuscript has been revised to address the comments is necessary.

We appreciate the opportunity to respond to your updated comments on our manuscript. We carefully considered the feedback and made changes to address the specific issues raised.

In response to your question of insufficient description of the data shown in Reviewer Fig. 1 and 2, we would like to confirm that we have taken this feedback seriously. Supplementary data, including the information provided in Reviewer Figures 1 and 2, have been fully described and integrated into the body of the manuscript according to your request. We ensured that the reliability and significance of new data were clearly presented to enhance the overall synthesis of the manuscript.

We are grateful to your valuable feedback, which undoubtedly contributed to the refinement of our manuscript. We hope that the revised version meets the standards of the journal and look forward to the opportunity for further deliberation.

Reviewer #2 (Recommendations For The Authors):

Additional feedback from the reviewer:

"I think the authors have been responsive to my previous comments. However, I cannot find this new data in the main text but rather only in the response to reviewers. New data should be incorporated into the main text not the supplement as the controls are important to consider alongside the treatment groups. Lastly, while the authors include BODIPY in their approaches, their results are not quantitative. My suggestion was to include this data in a quantitative manner not just the images. Lastly, I am still somewhat puzzled about the connection with GABA. The rationale for its selection other than it was significantly changed is not strong."

Thank you for providing us with the latest feedback. We appreciate the opportunity to address the specific concerns raised and provide a detailed response to each point.

(1) Incorporation of New Data into the Main Text:

We acknowledge the reviewer's comment regarding the incorporation of new data into the main text rather than solely in the response to reviewers. In response to this feedback, we have diligently revised the manuscript to ensure that the new data, including controls, is now seamlessly integrated into the main body of the text. This modification allows for a more comprehensive and contextual presentation of the data, as recommended by the reviewer.

(2) Quantitative Presentation of BODIPY Results:

We understand the importance of presenting quantitative data for the BODIPY results, and we appreciate the reviewer's suggestion to include this information in a quantitative manner, not just as images. In line with this valuable feedback, we have revised the relevant sections to incorporate quantitative data alongside the images, providing a more robust and comprehensive presentation of the results.

(3) Rationale for the Selection of GABA:

In the present study, in order to elucidate the molecular mechanisms through which pathway participates metformin-treated IR injury, we analysed gene expression profiles of each group mice, showing that similar mRNA changes are mainly concentrated in the three top pathways: lipid metabolism, carbohydrate metabolism, and amino acid metabolism. Given the close relevance between lipid metabolism and ferroptosis, and the fact of carbohydrate metabolism is a primary way to metabolize amino acids, 22 species of amino acid were detected in liver tissues using HPLC-MS/MS for further identification of key metabolites involved in the role of metformin against HIRI-induced ferroptosis. It was found that only GABA level is significantly increased by metformin treatment and FMT treatment, further verifying by the data of ELISA detection. Consequently, we identified GABA was the main metabolism of metformin protecting from HIRI and focus on the source of GABA generation.

We would like to express our gratitude to your thorough evaluation and constructive feedback, which has undoubtedly contributed to the improvement of our manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Hepatic ischemia/reperfusion injury (HIRI) is a common and inevitable factor leading to poor prognosis in various liver diseases, making the outcomes of current treatments in clinic unsatisfactory. Metformin has been demonstrated to be beneficial to alleviate HIRI in recent studies, however, the underpinning mechanism remains unclear. In this study, we found metformin mitigates HIRI-induced ferroptosis through reshaped gut microbiota in mice, which was confirmed by the results of fecal microbiota transplantation (FMT) treatment but showed the elimination of the beneficial effects when gut bacteria were depleted using antibiotics. Detailedly, through 16S rRNA and metagenomic sequencing, we identified that the metformin-reshaped microbiota was characterized by the increase of gamma-aminobutyric acid (GABA) producing bacteria. This increase was further confirmed by the elevation of GABA synthesis key enzymes, glutamic acid decarboxylase (GAD) and putrescine aminotransferase (PAT), in gut microbes of metformin-treated mice and healthy volunteers. Furthermore, the benefit of GABA against HIRI-induced ferroptosis was demonstrated in GABA-treated mice. Collectively, our data indicate that metformin can mitigate HIRI-induced ferroptosis by reshaped gut microbiota, with GABA identified as a key metabolite.

Introduction

Hepatic ischemia/reperfusion injury (HIRI) remains a major cause of liver damage mainly following liver transplantation, complicated hepatectomy, haemorrhagic shock and severe liver trauma.1–4 Traditional therapeutics have been developed and repurposed to treat HIRI, including ischemic preconditioning, machine perfusion technologies, gene interventions, and stem cell therapy.5–8 Unfortunately, current treatments are not satisfactory due to the complicated pathogenesis of HIRI. Recently, a novel theory relating to the gut-liver axis has been well accepted, demonstrating that gut microbiota is vital for liver diseases, however, there is no strategy targeted this theory to alleviate HIRI in clinical practice.

The liver links with the intestine by portal circulation which allows the transfer of gut-derived products, not only nutrients, but also microbial metabolites and components.9 As an important part of the gut-liver communication, the intestinal mucosal and vascular barriers limit the systemic dissemination of microbes and toxins while allowing nutrients to access the circulation and reaching the liver.10 Additionally, it is generally believed that liver injury can provoke intestinal mucosal damage and inflammation, further resulting in gut dysbiosis, which reversely induces a secondary attack on the liver.11 Microbial agents, such as synbiotics and probiotics, have demonstrated to attenuate different liver diseases in animal studies,12, 13 indicating that modulating gut microbes potentially prevents HIRI in clinical practice.

Recently, the hypoglycaemic drug, metformin, has attracted intensive attention in its pleiotropic biological actions in various diseases, such as obesity, cancer, nuroinflammation, especially liver diseases.14–17 As evidences shown, the gut with high concentration of metformin become an important action site to reshape microbiota.18 According to Sun et al.,19 metformin was proven to decrease the abundance of the B. fragilis species and to increase the levels of bile acid and glycoursodeoxycholic acid in the gut, which inhibited intestinal farnesoid X receptor against obesity-related metabolic dysfunction in mice. Moreover, it is also demonstrated that metformin could attenuate sepsis-related liver injury by increasing the proportion of Bifidobacterium, Muribaculaceae, Parabacteroides distasonis, and Alloprevitella in rats.20 As for the HIRI, metformin has also shown the protective effect,21 but the detailed role of gut microbiota remains unclear.

Although metformin showed protective effects on HIRI, the role and distinctions of reshaping microbes have yet to be determined. In the present study, we sought to elucidate that gamma- aminobutyric acid (GABA)-producing bacteria played a critical role for the protective effects of metformin against HIRI by using 16S rRNA and metagenomic sequencing to find the different features of gut microbiota. The finding was reconfirmed by the results of mice that were treated with metformin-reshaped fecal microbiota or oral antibiotics administration. Furthermore, GABA was found to show a protective effect on HIRI by intraperitoneally injection into mice in advance, while an increase of microbial GABA was observed in both mice and human experiments after metformin treatment. Evidenced by the improvement of HIRI, we hypothesize that metformin-reshaped gut microbiota elicits greater therapeutic responses for liver diseases and serve as a promising targeted therapeutic.

Results

Metformin significantly mitigates HIRI and reshapes gut microbiota in mice

In order to evaluate the protective effect of metformin on HIRI, survival experiment was employed in mice. The survival rate of HIRI mice increased from 41.67% to 90.00% after metformin intervention (Fig. 1a), accompanied by hepatic morphological and histopathological improvement including the recovery of liver pathological structure, the decrease in red blood cell aggregation and consumption of liver glycogen (Fig. 1b), as well as the decrease in serum activity of ALT and AST (Fig. 1c). Qualification of MDA and GSH concentrations in the mice liver demonstrated that the presence of metformin significantly attenuated HIRI-induced oxidation (Fig. 1d).

Removed in the revised version: Metformin significantly mitigates HIRI by reshaping gut microbiota in mice

Earlier version: (a) Survival rate analysis (n = 10-12/group), Kaplan Meier was used to analyze statistical differences (Control vs IR, and IR+Met vs IR). (b) Morphological appearance, H&E and PAS staining of liver (n = 3/group). (c) Serum ALT and AST levels. (d) Levels of MDA and GSH in liver tissue. (e) Sobs index, PCoA (n = 4/group). (f) H&E staining of ileum and colon (n = 3/group). (g) Western blot of Occludin in colon (n = 3/group). (h) IFC of ZO-1 and Occludin in colon (n = 3/group). (i) qRT-PCR analysis of Occludin, ZO-1, Claudin-1, JAM1 and JAM4 in ileum. (j) IL-6, IL-1β and IL-18 in ileum and colon. (n = 3–5/group). (k) FD-4 level in mice serum (n = 3). Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Revised version: Metformin significantly mitigates HIRI by reshaping gut microbiota in mice (a) Survival rate analysis (n = 10-12/group), Kaplan Meier was used to analyze statistical differences (Control vs IR, and IR+Met vs IR). (b) Morphological appearance, H&E and PAS staining of liver (n = 3/group). (c) Serum ALT and AST levels. (d) Levels of MDA and GSH in liver tissue. (e) Sobs index, PCoA (n = 4/group). (f) H&E staining of ileum and colon (n = 3/group). (g) Western blot of Occludin in colon (n = 3/group). (h) IFC of ZO-1 and Occludin in colon (n = 3/group). (i) qRT-PCR analysis of Ocln, Tjp1, Cldn7, F11r and Igsf5 in ileum. (j) Il6, Il1b and Il18 in ileum and colon. (n = 3–5/group). (k) FD-4 level in mice serum (n = 3). Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Earlier version: In order to clarify the critical role of gut microbiota in the pleiotropic actions of metformin,22–24 fecal samples were collected from the mice to perform 16S rRNA sequencing. It was found that the gut microbiota was substantially modulated in IR + metformin (Met) group, showing decreased abundance (alpha diversity) and reshaped distribution (Fig. 1e). Next, as the microbiota-gut-liver axis theory indicates,25 HIRI-induced dysfunction of the gut barrier may aggravate liver damage by disrupting the gut microbiota. H&E staining of the colon and ileum, mainly involved in gut barrier dysfunction, showed that the density and integrity of intestinal mucosa were strikingly reduced by HIRI, which was restored by metformin treatment (Fig. 1f). The immunofluorescence analysis of colon tissues showed that the decreased tight junction protein ZO-1 and Occludin in HIRI-treated mice were reversed by metformin (Fig. 1g and 1h). Similarly, western blot analysis of Occludin in colon (Fig. 1g) and the mRNA levels of Occludin, ZO-1, Claudin-1, JAM1 and JAM4 in ileum of HIRI mice were increased with metformin administration (Fig. 1i). The inflammatory cytokine IL-6, IL-1β and IL-18 were significantly upregulated in the ileum and colon received HIRI and were reduced by metformin treatment (Fig. 1j). In addition, the serum FD-4 level in IR group was cut down by metformin (Fig. 1k). These results suggested that metformin plays a critical role in restoring gut barrier function and alleviating HIRI.

Revised version: In order to clarify the critical role of gut microbiota in the pleiotropic actions of metformin,22–24 fecal samples were collected from the mice to perform 16S rRNA sequencing. It was found that the gut microbiota was substantially modulated in IR + metformin (Met) group, showing decreased abundance (alpha diversity) and reshaped distribution (Fig. 1e). Next, as the microbiota-gut-liver axis theory indicates,25 HIRI-induced dysfunction of the gut barrier may aggravate liver damage by disrupting the gut microbiota. H&E staining of the colon and ileum, mainly involved in gut barrier dysfunction, showed that the density and integrity of intestinal mucosa were strikingly reduced by HIRI, which was restored by metformin treatment (Fig. 1f). The immunofluorescence analysis of colon tissues showed that the decreased tight junction protein ZO-1 and Occludin in HIRI-treated mice were reversed by metformin (Fig. 1g and 1h). Similarly, western blot analysis of Occludin in colon (Fig. 1g) and the mRNA levels of Ocln, Tjp1, Cldn7, F11r and Igsf5 in ileum of HIRI mice were increased with metformin administration (Fig. 1i). The inflammatory cytokine Il6, Il1b and Il18 were significantly upregulated in the ileum and colon received HIRI and were reduced by metformin treatment (Fig. 1j). In addition, the serum FD-4 level in IR group was cut down by metformin (Fig. 1k). These results suggested that metformin plays a critical role in restoring gut barrier function and alleviating HIRI.

Metformin-reshaped fecal microbiota attenuates HIRI

Earlier version: In order to further confirm the role of metformin in reshaping the gut microbiota, microbes were transplanted into mice before inducing reperfusion (IR). As expected, the survival analysis of mice in the IR + fecal microbiota transplantation (FMT) group showed a similarly protective efficacy as the IR + Met group, as indicated by the significantly improved survival rate of IR mice (Fig. 2a). This finding was validated by ALT and AST levels as well as H&E, PAS and DHE staining (Fig. 2b and 2e). FMT treatment could restore the GSH level but reduce the accumulation of MDA in hepatic tissues (Fig 2c and 2d). Furthermore, a decreased alpha diversity was observed in the FMT group, similar to IR + Met group (Fig. 2f). Meanwhile, the attenuation of gut damage was showed by H&E staining, and the improvement in tight junction proteins ZO-1 and Occludin, and decreased HIRI-induced IL-1β, IL-6, and IL-18 (Fig. 2g-2i). Serum FD-4 level showed a sharply decrease in IR + FMT group (Fig. 2j). To further identify the function of gut microbes, experiments were designed, and combination treatment of antibiotics (1 mg/mL penicillin sulfate, 1 mg/mL neomycin sulfate, 1 mg/mL metronidazole and 0.16 mg/mL gentamicin) and metformin were employed for 1 week before IR treated. As results showed, the survival rate of IR + Met + antibiotics (Abx) group was greatly reduced, while the levels of ALT, AST, MAD and GSH were completely opposite to IR + Met group (Fig. 2k). The results confirmed that gut microbiota is critical for the effect of metformin against HIRI.

Revised version: In order to further confirm the role of metformin in reshaping the gut microbiota, microbes were transplanted into mice before inducing ischemia reperfusion (IR). As expected, the survival analysis of mice in the IR + fecal microbiota transplantation (FMT) group showed a similarly protective efficacy as the IR + Met group, as indicated by the significantly improved survival rate of IR mice (Fig. 2a). This finding was validated by ALT and AST levels as well as H&E, PAS and DHE staining (Fig. 2b and 2e). FMT treatment could restore the GSH level but reduce the accumulation of MDA in hepatic tissues (Fig 2c and 2d). Furthermore, a decreased alpha diversity was observed in the FMT group, similar to IR + Met group (Fig. 2f). Meanwhile, the attenuation of gut damage was showed by H&E staining, and the improvement in tight junction proteins ZO-1 and Occludin, and decreased HIRI-induced Il1b, Il6, and Il18 (Fig. 2g-2i). Serum FD-4 level showed a sharply decrease in IR + FMT group (Fig. 2j). To further identify the function of gut microbes, experiments were designed, and combination treatment of antibiotics (1 mg/mL penicillin sulfate, 1 mg/mL neomycin sulfate, 1 mg/mL metronidazole and 0.16 mg/mL gentamicin) and metformin were employed for 1 week before IR treated. As results showed, the ALT, AST and MDA detection in control group treated with Abx showed no obvious change to control group, indicating that the Abx had no influence on the liver. The survival rate of IR + Met + antibiotics (Abx) group was greatly reduced, while the levels of ALT, AST, MAD and GSH were completely opposite to IR + Met group (Fig. 2k). The results confirmed that gut microbiota is critical for the effect of metformin against HIRI.

Removed in the revised version: Metformin-reshaped fecal microbiota attenuates HIRI

Earlier version: (a) Survival rate analysis (n = 10-12/group), Kaplan Meier was used to analyze statistical differences. (b) Serum ALT and AST levels. Levels of MDA (c) and GSH (d) in liver tissue. (e) H&E, PAS and DHE staining of liver (n = 3/group). (f) Alpha diversity and PCoA level of mice microbes. (g) H&E staining of ileum and colon and Western blot analysis of Occludin in colon (n = 3/group). (h) IFC of ZO-1 and Occludin in colon (n = 3/group). (i) qRT-PCR analysis of Occludin, ZO-1, Claudin-1, JAM1 and JAM4 in ileum, and IL-6, IL-1β and IL-18 in ileum and colon. (n = 3–5/group). (j) FD-4 level in mice serum (n = 3). (k) Survival rate of IR + Met + Abx group (n = 10), Kaplan Meier was used to analyze statistical differences and the levels of serum ALT, AST, liver MDA, Fe, and GSH in IR + Met + Abx group. Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Revised version: Metformin-reshaped fecal microbiota attenuates HIRI (a) Survival rate analysis (n = 10-12/group), Kaplan Meier was used to analyze statistical differences. (b) Serum ALT and AST levels. Levels of MDA (c) and GSH (d) in liver tissue. (e) H&E, PAS and DHE staining of liver (n = 3/group). (f) Alpha diversity and PCoA level of mice microbes. (g) H&E staining of ileum and colon and Western blot analysis of Occludin in colon (n = 3/group). (h) IFC of ZO-1 and Occludin in colon (n = 3/group). (i) qRT-PCR analysis of Ocln, Tjp1, Cldn7, F11r and Igsf5 in ileum, and Il6, Il1b and Il18 in ileum and colon. (n = 3– 5/group). (j) FD-4 level in mice serum (n = 3). (k) ALT, AST and MDA detection between Control and Control + Abx group; survival rate of IR + Met + Abx group (n = 10), Kaplan Meier was used to analyze statistical differences ; and the levels of serum ALT, AST, liver MDA, Fe, and GSH in IR + Met + Abx group. Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

FMT alleviates HIRI-induced ferroptosis through reshaped fecal microbiota

Earlier version: Since the reperfusion-induced excessive oxidative stress takes place in iron-rich hepatocytes,26 experiments were designed to observe ferroptosis in HIRI. The ferroptosis inhibitor, DFO, significantly ameliorated the liver injury, suggesting that ferroptosis contributes to HIRI (sFig. 1). Hence, to find out the relation, ferroptosis-related indicators were also monitored in both IR + Met and IR + FMT groups. The increased accumulation of Fe after HIRI was significantly cut down by metformin and FMT treatment (Fig. 3a). The quantitative Real-time PCR (qRT-PCR) analysis showed that the HIRI-induced upregulation of ferroptotic parameters, ACSL4, SLC7A11, SLC39A14 and COX-2, were significantly decreased by metformin and its reshaped gut microbiota, respectively (Fig. 3b-3e). The results of western blotting further confirmed the anti-ferroptosis of metformin and FMT by the remarkable reduction of ACSL4, TFR1, VDAC1, VDAC2 and VDAC3 levels and the apparent increase of antiferroptotic FTH1 and Xct levels (Fig. 3f), The results was further consolidated by lipid ROS staining and transmission electron microscopy (TEM) analysis (Fig. 3g). Additionally, transcriptome data on genes associated with antiferroptosis showed an increase in both IR + Met and IR + FMT groups, whereas proferroptic genes showed a decrease, when compared with IR group (Fig. 3h). Collectively, the results suggested that FMT performed protection against HIRI by inhibiting ferroptosis.

Revised version: Since the reperfusion-induced excessive oxidative stress takes place in iron-rich hepatocytes,26 experiments were designed to observe ferroptosis in HIRI. The ferroptosis inhibitor, DFO, significantly ameliorated the liver injury, suggesting that ferroptosis contributes to HIRI (sFig. 1). Hence, to find out the relation, ferroptosis-related indicators were also monitored in both IR + Met and IR + FMT groups. The increased accumulation of Fe after HIRI was significantly cut down by metformin and FMT treatment (Fig. 3a). The quantitative Real-time PCR (qRT-PCR) analysis showed that the HIRI-induced upregulation of ferroptotic parameters, Acsl4, Slc7a11, Slc39a14 and Ptgs2, were significantly decreased by metformin and its reshaped gut microbiota, respectively (Fig. 3b-3e). The results of western blotting further confirmed the anti-ferroptosis of metformin and FMT by the remarkable reduction of ACSL4, TFR1, VDAC1, VDAC2 and VDAC3 levels and the apparent increase of antiferroptotic FTH1 and Xct levels (Fig. 3f), The results were further consolidated by lipid ROS staining and transmission electron microscopy (TEM) analysis (Fig. 3g). Additionally, transcriptome data on genes associated with antiferroptosis showed an increase in both IR + Met and IR + FMT groups, whereas proferroptic genes showed a decrease, when compared with IR group (Fig. 3h). Collectively, the results suggested that FMT performed protection against HIRI by inhibiting ferroptosis.

Removed in the revised version: Metformin alleviates HIRI through inhibiting ferroptosis

Earlier version: (a) Levels of Fe in liver tissue (n=6). qRT-PCR analysis of ACSL4 (b) SLC7A11 (c), SLC39A14 (d) and COX-2 (e) (n = 5-8/group). (f) Western blot analysis of ACSL4, FTH1, VDAC1, VDAC2, VDAC3, TFR1 and Xct in liver tissue (n = 3/group). (g) Lipid ROS staining and TEM analysis of liver tissues (n = 3/group). (h) Ferroptosis-related transcriptome cluster diagram heatmap. Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Revised version: Metformin alleviates HIRI through inhibiting ferroptosis (a) Levels of Fe in liver tissue (n=6). qRT-PCR analysis of Acsl4 (b) Slc7a11 (c), Slc39a14 (d) and Ptgs2 (e) (n = 5-8/group). (f) Western blot analysis of ACSL4, FTH1, VDAC1, VDAC2, VDAC3, TFR1 and Xct in liver tissue (n = 3/group). (g) Lipid ROS staining and TEM analysis of liver tissues (n = 3/group). (h) Ferroptosis-related transcriptome cluster diagram heatmap. Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Metformin induces GABA-producing gut microbiota

Earlier version: To elucidate the molecular mechanisms through which pathway participates metformin-treated IR injury, we analysed gene expression profiles of each group mice. Transcriptome sequencing analysis revealed that 9697 genes were in common among four groups (sFig. 6). Therefore, we used these common genes for KEGG analysis, showing that similar mRNA changes are mainly concentrated in the three top pathways: lipid metabolism, carbohydrate metabolism, and amino acid metabolism (Fig 4a). Given the close relevance between lipid metabolism and ferroptosis (sFig. 2), and the fact of carbohydrate metabolism is a primary way to metabolize amino acids, 22 species of amino acid were detected in liver tissues using HPLC-MS/MS for further identification of key metabolites involved in the role of metformin against HIRI-induced ferroptosis. It was found that only GABA level is significantly increased by metformin treatment and FMT treatment (sFig. 3), further verifying by the data of ELISA detection (Fig. 4b). Moreover, it was observed that the genus of Bacteroides had a significant increase based on the 16s rRNA gene sequencing of metformin-treated mice microbes (Fig 4c). qRT-PCR analysis of eukaryotic GABA synthesis key enzymes including SSADH, GABA-T, Gad1 and Gad2 in liver tissues, were not upregulated by metformin (Fig. 4d), which was also confirmed by western blotting detection of GAD1 and GAD2 in liver (Fig. 4f). Interestingly, the fecal GABA concentration was shown a significant increase in metformin-treated mice, which was copied by the IR + FMT group, but eliminated by antibiotics treatment (Fig. 4e). In order to further identify the increased GABA originates from gut microbiota, two key enzymes of prokaryotes GABA synthesis, GAD and PAT, were detected on DNA level, finding that both of them are significantly increased in the feces from IR+Met and IR+FMT groups (Fig. 4g and h). In order to clarify the specific effects of metformin on microbiota, given the big safety margin, healthy volunteers were recruited for a 1 week of daily oral 500mg dose of metformin trial. Fecal samples were collected before and after oral administration of metformin for metagenomic analysis. Detailedly, the species of Bacteroides containing Bacteroides thetaiotaomicron, Bacteroides unifomis, and Bacteroides salyersiae, were enriched in human gut after metformin administration (Fig. 4i). Excitingly, the increase of GABA production, and the expression of GAD and PAT were shown in the results of fecal experiments in vitro and in vivo, accompanied by the decreased glutamic acid (Glu), a key substrate for gut microbiota to synthesize GABA (Fig. 4j and k). Besides, putrescine, another substrate of GABA synthesis, was not found in primary culture medium but appeared after 24h fermentation and then was consumption later in the anaerobic culture of human feces experiment. (Fig. 4l). Taken collectively, GABA- producing gut microbiota was upregulated by metformin.

Revised version: To elucidate the molecular mechanisms through which pathway participates metformin-treated IR injury, we analysed gene expression profiles of each group mice. Transcriptome sequencing analysis revealed that 9697 genes were in common among four groups (sfig. 6). Therefore, we used these common genes for KEGG analysis, showing that similar mRNA changes are mainly concentrated in the three top pathways: lipid metabolism, carbohydrate metabolism, and amino acid metabolism (Fig 4a). Given the close relevance between lipid metabolism and ferroptosis (sFig. 2), and the fact of carbohydrate metabolism is a primary way to metabolize amino acids, 22 species of amino acid were detected in liver tissues using HPLC-MS/MS for further identification of key metabolites involved in the role of metformin against HIRI-induced ferroptosis. It was found that only GABA level is significantly increased by metformin treatment and FMT treatment (sFig. 3), further verifying by the data of ELISA detection (Fig. 4b). Moreover, it was observed that the genus of Bacteroides had a significant increase based on the 16s rRNA gene sequencing of metformin-treated mice microbes. qRT-PCR analysis of eukaryotic GABA synthesis key enzymes including Aldh5a1, Abat, Gad1 and Gad2 in liver tissues, were not upregulated by metformin (Fig. 4d), which was also confirmed by western blotting detection of GAD1 and GAD2 in liver (Fig. 4f). Interestingly, the fecal GABA concentration was shown a significant increase in metformin-treated mice, which was copied by the IR + FMT group, but eliminated by antibiotics treatment (Fig. 4e). In order to further identify the increased GABA originates from gut microbiota, two key enzymes of prokaryotes GABA synthesis, GAD and PAT, were detected on DNA level, finding that both of them are significantly increased in the feces from IR+Met and IR+FMT groups (Fig. 4g and h). In order to clarify the specific effects of metformin on microbiota, given the big safety margin, healthy volunteers were recruited for a 1 week of daily oral 500mg dose of metformin trial. Fecal samples were collected before and after oral administration of metformin for metagenomic analysis. Detailedly, the species of Bacteroides containing Bacteroides thetaiotaomicron, Bacteroides unifomis, and Bacteroides salyersiae, were enriched in human gut after metformin administration (Fig. 4i). Excitingly, the increase of GABA production, and the expression of GAD and PAT were shown in the results of fecal experiments in vitro and in vivo, accompanied by the decreased glutamic acid (Glu), a key substrate for gut microbiota to synthesize GABA (Fig. 4j and k). Besides, putrescine, another substrate of GABA synthesis, was not found in primary culture medium but appeared after 24h fermentation and then was consumption later in the anaerobic culture of human feces experiment. (Fig. 4l). Taken collectively, GABA- producing gut microbiota was upregulated by metformin.

Removed in the revised version: Metformin induces GABA-producing gut microbiota.

Earlier version: (a) KEGG analysis of transcriptome. (b). HPLC-MS/MS of GABA content in liver, ELISA of GABA in liver (n = 6–9/group). (c) Community heatmap analysis on genus level. (d) qRT-PCR analysis of key metabolite enzyme glutamate decarboxylase 1 (Gad1), glutamic acid decarboxylase 2 (Gad2), 4-aminobutyrate aminotransferase (Gabat), aldhehyde dehydrogenase family 5 (SSADH) in liver (n = 5-8/group). (e) ELISA of GABA in feces (n = 7–10/group) (f) Western-blot analysis of GAD1 and GAD2 in liver. (g-h) qRT-PCR analysis of GAD and PAT in mice feces (n= 3-5/group). (i) Metagenomic sequencing analysis of human feces. (j-k) The level of GAD, PAT, GABA and Glu in human feces (n= 5-6/group) l) The level of GABA, Glu and putrescine in human feces vitro culture (n= 3/group). Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Revised version: Metformin induces GABA-producing gut microbiota. (a) KEGG analysis of transcriptome. (b). HPLC-MS/MS of GABA content in liver, ELISA of GABA in liver (n = 6–9/group). (c) Community heatmap analysis on genus level. (d) qRT-PCR analysis of key metabolite enzyme glutamate decarboxylase 1 (Gad1), glutamic acid decarboxylase 2 (Gad2), 4-aminobutyrate aminotransferase (Abat), aldhehyde dehydrogenase family 5 (Aldh5a1) in liver (n = 5-8/group). (e) ELISA of GABA in feces (n = 7–10/group) (f) Western-blot analysis of GAD1 and GAD2 in liver. (g-h) qRT-PCR analysis of GAD and PAT in mice feces (n= 3-5/group). (i) Metagenomic sequencing analysis of human feces. (j-k) The level of GAD, PAT, GABA and Glu in human feces (n= 5-6/group) l) The level of GABA, Glu and putrescine in human feces vitro culture (n= 3/group). Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

GABA is the critical metabolite of metformin-reshaped gut microbiota against HIRI-induced ferroptosis

The animal experiment was designed to further observe the effects of GABA on HIRI-induced ferroptosis (Fig. 5a). As expected, the data showed that GABA treatment reduced histological damage and liver glycogen consumption (Fig. 5b). Serum liver function detection showed that ALT and AST activities were dramatically depressed after GABA treatment (Fig. 5c). The decreased MDA and Fe as well as increased GSH were observed in further determination, confirming the protection of GABA on liver (Fig. 5d). The results of western blotting showed the anti-ferroptosis effect of GABA treatment by the significant reduction in ACSL4, TFR1, VDAC1,2,3 and increase of anti-ferroptotic FTH1 levels, which had the same attenuating effect as metformin and FMT treatment (Fig. 5e). In addition, the clustergram showed GABA upregulated the genes of antiferroptosis and depressed the proferroptic ones, which was similar to the effect after metformin treatment (Fig. 5f). In summary, GABA treatment reduced HIRI- induced ferroptosis.

Removed in the revised version: GABA is the critical metabolite of metformin-reshaped gut microbiota against HIRI-induced ferroptosis.

Earlier version: (a) Diagram of GABA treatment experiment on mice model. (b) H&E and PAS staining of liver (n = 3/group). (c) Serum ALT and AST levels (n =6/group). Levels of MDA, GSH and Fe in liver tissues (d) (n =4/group). (e) Western blot analysis of ASCL4, FTH1, TFR1 and VADC1,2,3 in liver (n = 4/group). (f) Ferroptosis-related transcriptome cluster diagram heatmap (n = 3-5/group). Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Revised version: GABA is the critical metabolite of metformin-reshaped gut microbiota against HIRI-induced ferroptosis. (a) Diagram of GABA treatment experiment on mice model. (b) H&E and PAS staining of liver (n = 3/group). (c) Serum ALT and AST levels (n =6/group). Levels of MDA, GSH and Fe in liver tissues (d) (n =4/group). (e) Western blot analysis of ASCL4, FTH1, TFR1 and VADC1,2,3 in liver (n = 4/group). (f) Ferroptosis-related transcriptome cluster diagram heatmap (n = 3-5/group). Data are expressed as mean ± standard deviation. One-way ANOVA was used to analyze statistical differences; *P < 0.05, **P < 0.01 and ***P < 0.001.

Discussion

Metformin has been well researched for treating various liver diseases due to its multifarious effects. Notably, it has been well documented that oral administration of metformin lead to therapeutic effects on HIRI,27 leaving a thought-provoking question in its underlying mechanism. In the current study, the changes in the gut microbiota altered by metformin are demonstrated to be necessary and sufficient for conferring liver injury protection. Moreover, it has been confirmed that metformin attenuated HIRI by the increased GABA from reshaped gut microbiota, through anti-ferroptosis.

Increasing evidence shows that the gut microbiota is critical for the pleiotropic actions of metformin.25, 28, 29 Interestingly, metformin was previously found to have a much higher concentration in the gut lumen than other organs,18 as also confirmed by the current study. In fact, that concentration favoured the action of metformin on reshaping gut microbes. Despite the distinct heterogeneity of metformin-reshaped gut microbiota in recent research, Bacteroides with a high abundance of gut microbes was reported to be reduced after metformin treatment in studies of hyperlipidemia and diabetes,19, 30 which other studies vary from for the different action of this bacteria.31, 32 In our study, we used a disease model of HIRI, which may have unique characteristics compared to other disease models. It is possible that the specific disease model influenced the response of the gut microbiota. Additionally, the starting microbiota of the recipients and the characteristics of the donor microbiota used for FMT could also play a role in the disparity. Our findings showed an increase abundance of the genus of Bacteroides after metformin treatment supported by the data of mice and healthy volunteers. Detailedly, the increased Bacteroides especially Bacteroides thetaiotaomicron, Bacteroides unifomis, and Bacteroides salyersiae were observed in the human gut after metformin administration. The enriched Bacteroides theaiotaomicron, a glutamate-fermenting commensal, was reported to be decreased in obese individuals and was inversely correlated with serum glutamate concentration.33 Therefore, Bacteroides theaiotaomicron might be the potential bacterium regulating of GABA production. Little research has been done on Bacteroides unifomis and Bacteroides salyersiae, needing a further exploration of their characteristics in the future. Notably, Phascolarctobacterium, coexistly with Bacteroides thetaiotaomicron,34 was reported as an emerging probiotic which showed a significant increase in our study. Taken collectively, those reshaped bacteria after metformin treatment may be responsible for attenuating HIRI.

Earlier version: Increased GABA was shown from the metformin reshaped gut microbiota for the first time, verified by our results of feces from mice and humans. As the critical metabolite of bacteria against acid stress, GABA was extensively produced by gut microbes, especially the species with strong capacity of acid resistance. Interestingly, our analysis of the sequences in PUBMED Genebank found that those three species of Bacteroides, Bacteroides thetaiotaomicron, Bacteroides unifomis, and Bacteroides salyersiae, which belong to metformin reshaped microbiota, contain the key enzyme genes of GABA synthesis GAD and PAT (sFig 4), further confirmed by the qRT-PCR result of human fecal DNA in vivo and in vitro experiment. Specifically, the species with the strong ability of producing GABA such as Bacteroides theaiotaomicron, Bacteroides unifomis and Bacteroides salyersiae were increased in metformin-treated volunteers. Some studies were published to elucidate the mechanism of metformin regulating bile acids and folate metabolite by reshaped gut microbiota, however, there is no report about how metformin to affect GABA synthesis of gut microbes. The data from anaerobic cultured human feces showed that metformin promoted the increase of GAD and PAT, and GABA level along with the substrate consumption of glutamate and putrescine in a time-dependent manner.

Revised version: Increased GABA was shown from the metformin reshaped gut microbiota for the first time, verified by our results of feces from mice and humans. As the critical metabolite of bacteria against acid stress, GABA was extensively produced by gut microbes, especially the species with strong capacity of acid resistance. Interestingly, our analysis of the sequences in PUBMED Gene bank found that those three species of Bacteroides, Bacteroides thetaiotaomicron, Bacteroides unifomis, and Bacteroides salyersiae, which belong to metformin reshaped microbiota, contain the key enzyme genes of GABA synthesis GAD and PAT (sFig 4), further confirmed by the qRT-PCR result of human fecal DNA in vivo and in vitro experiment. Specifically, the species with the strong ability of producing GABA such as Bacteroides theaiotaomicron, Bacteroides unifomis and Bacteroides salyersiae were increased in metformin-treated volunteers. Some studies were published to elucidate the mechanism of metformin regulating bile acids and folate metabolite by reshaped gut microbiota, however, there is no report about how metformin to affect GABA synthesis of gut microbes. The data from anaerobic cultured human feces showed that metformin promoted the increase of GAD and PAT, and GABA level along with the substrate consumption of glutamate and putrescine in a time-dependent manner.

GABA production was reported to be critical against acid stress for microbiota to survive and preserve metabolic activity, especially in the human gut with a lower pH than other animals.35 Humans benefit from bacteria-derived GABA. Peripheral GABA was first demonstrated in intestinal flora in 1981 and further confirmed by recent studies.36 In addition to the important role in the central neural system and immune system, GABA showed a protective effect in different acute liver injury models.37, 38 Published papers have shown that GABA performs beneficial effect through binding to the GABA receptors or promoting histone acetylation.39, 40 Our data from GABA-treated mice showed protective effects against HIRI-induced ferroptosis. However, further studies are required to examine the exact molecular mechanisms of GABA regulating ferroptosis.

In conclusion, we have demonstrated metformin-reshaped gut microbiota performs therapeutic responses for HIRI. Gut microbiota has been obviously modified while receiving oral metformin treatment, suggesting that metformin-reshaped gut microbiota has enhanced intestinal mucosal function so that reduce harmful substances to the liver, which most efficiently evade damage. Specifically, GABA producing bacteria were significantly regulated after metformin administration. Demonstrated for the first time, we showed that the increased GABA is a key metabolite of gut microbiota for metformin to mitigate HIRI. These results offer a unique therapeutic strategy, with optimized gut microbiota, that can serve as a target for metformin treating liver diseases.

Material and methods

Animals

Six-to-eight-week-old specific pathogen-free male C57BL/6 mice were purchased from Beijing Weitonglihua Experimental Animal Technology Co. Ltd. (Beijing, China). All mice were housed under controlled temperature and humidity conditions with a 12-h light–dark cycle and free access to food and water. The mice were fasted overnight before the experiments. All animal experiments were approved by the Institutional Animal Care and Use Committee of the Experimental Animal Centre of Wenzhou Medical University (wydw2021-0348).

Healthy volunteers

All research was conducted in accordance with both the Declarations of Helsinki and Istanbul, all research was approved by the Ethics Committee of the Second Affiliated Hospital of Wenzhou Medical University (2022-K-08-01), and written consent was given in writing by all subjects.

Hepatic ischemia/reperfusion (IR) mouse model

A mouse model for HIRI was established as previously described.3 Briefly, mice were anaesthetised with pentobarbital sodium. A noninvasive microvascular artery clip was placed on the left branches of the hepatic artery and portal vein for 30 min, and the clip was removed for reperfusion for 6 h. During the study period, body temperature was maintained at 37 °C using a heating lamp.

Experimental design

Earlier version: Mice were randomly divided into the several groups. In general, the arterial and portal vessels to the cephalad lobes were clamped for 30min and reperfusion for 6 h. No vascular occlusion in sham-controlled mice. In the treatment groups, animals were treated with metformin (1 mg/mL), metformin modified fecal microbiota transplantation or metformin (1 mg/mL) mixed with antibiotics (1 mg/mL penicillin sulfate, 1 mg/mL neomycin sulfate, 1 mg/mL metronidazole and 0.16 mg/mL gentamicin), in the drinking water for 1 week. Specifically, FMT-treated mice was fed with washed metformin-treated mice feces, which had been detected by HPLC showing that there was no metformin left in the feces. Other two groups of mice were infused at 1 h prior to the onset of liver ischemia with a single dose of deferoxamine (DFO) (20 mg/kg i.p.) or GABA (100 mg/kg i.p.), dissolved in PBS. Mice were sacrificed after reperfusion; liver and serum samples were collected for analysis.

Revised version: Mice were randomly divided into the several groups. In general, the arterial and portal vessels to the cephalad lobes were clamped for 30min and reperfusion for 6 h. No vascular occlusion in sham-controlled mice. In the treatment groups, animals were treated with metformin (1 mg/mL), metformin modified fecal microbiota transplantation or metformin (1 mg/mL) mixed with antibiotics (1 mg/mL penicillin, 1 mg/mL neomycin sulfate, 1 mg/mL metronidazole and 0.16 mg/mL gentamicin), in the drinking water for 1 week. Specifically, FMT-treated mice were fed with washed metformin-treated mice feces, which had been detected by HPLC showing that there was no metformin left in the feces. Other two groups of mice were infused at 1 h prior to the onset of liver ischemia with a single dose of deferoxamine (DFO) (20 mg/kg i.p.) or GABA (100 mg/kg i.p.), dissolved in PBS. Mice were sacrificed after reperfusion; liver and serum samples were collected for analysis.

Measurement of serum alanine transaminase (AST), aspertate aminotransferase (ALT) level, liver Fe content and GABA level

Serum AST and ALT levels were measured using assay kits (Nanjing Jiancheng Bioengineering Institute, C010-2-1, C009-2-1). Liver samples were homogenised in saline, and Fe concentrations were measured by tissue Fe assay kit (Nanjing Jiancheng Bioengineering Institute, A039-2-1). GABA level was detected by an ELISA kit (Wuhan MSKBIO, KT21124).

Measurement of lipid peroxidation, glutathione (GSH) and superoxide anion levels of liver tissues

Malondialdehyde (MDA) Detection Kit (Nanjing Jiancheng Bioengineering Institute, A003-1-2) was determined to select the MDA level as a marker of lipid peroxidation. The assay was performed according to the manufacturer’s instructions. Lipid reactive oxygen species (ROS) was measured by BODIPY 581/591 C11 (MKbio, 217075-36-0). Sections were stained with BODIPY 581/591 C11 and nucleated with 4’,6-diamidino-2-phenylindole (DAPI) (Beyotime Institute of Biotechnology, A0131). Images were acquired under a fluorescence microscope (Nikon, Tokyo, Japan). The liver was homogenised, and the supernatant was collected for GSH analysis using a GSH assay kit (Nanjing Jiancheng Bioengineering Institute, A006). Frozen sections of the liver (8Lμm) were placed on glass slides and incubated with 10Lmmol/L dihydroethidium (DHE) (Beyotime Institute of Biotechnology, S0063) in a dark container at 37 °C for 30Lmin. Sections were rinsed three times in phosphate-buffered saline (PBS) then observed using an upright microscope (Nikon, Tokyo, Japan).

Hematoxylin-eosin (H&E) and periodic acid-schiff (PAS) staining

For histopathological examination, liver and intestinal tissues were fixed in 4% paraformaldehyde overnight. Selected tissue blocks were processed using a routine overnight cycle in a tissue processor. Tissue blocks were then embedded in wax and serially sliced into 5- μm sections. H&E and PAS staining performed as instructions provided by the reagent manufacturer (Solarbio, G1120, G1281). Visualization and images were obtained under an optical microscope (Nikon, Tokyo, Japan).

Histoimmunofluorescence staining

During immunofluorescence, sections were dyeing with primary antibody (1:100) Occludin (Proteintech, 27260-1-AP) and zonula occludens 1 (ZO-1) (Proteintech, 66378-1-IG) and incubated at 4°C overnight. Secondary antibody coupled with Alexa Fluor 488 (Bioss, bs-40296G-AF488) was incubated at 37°C for 1 h. After additional PBS washing, the sections were sealed with an antifluorescence quencher containing 4’,6-diamidino-2-phenylindole (DAPI). Immunofluorescence images were obtained using a laser scanning confocal microscope (Nikon, Tokyo, Japan) and quantified by ImageJ (NIH, USA).

Transmission electron microscopy (TEM) analysis

After reperfusion, the mice were euthanized, and the livers were excised and washed with precooled PBS (pH 7.4). Part of the liver was then removed and incubated overnight in 0.1LM PBS (pH 7.4) containing 2.5% glutaraldehyde. The target tissues were cut into 50-µm-thick sections using a vibratome. Selected areas of the livers were postfixed in 1% osmium tetroxide for 1Lh, dehydrated in a graded ethanol series and embedded in epoxy resin. The polymerisation was performed at 80L°C for 24Lh. Ultrathin sections (100Lnm) were cut, stained with uranyl acetate and lead citrate and viewed under a JEM2000EX TEM (JEOL, Tokyo, Japan). Five fields were randomly selected for each sample to examine mitochondria with ferroptosis features.

Quantitative Real-time PCR (qRT-PCR) analysis

Earlier version: Total RNA was extracted from liver and intestinal tissues using TRIzol (Yamei, YY101). Isolated RNA was reverse transcribed to cDNA using a kit (Vazyme, R323-01). The obtained cDNA was subjected to PCR using primers designed to detect Interleukin (IL)-6, IL-1β, IL-18, recombinant Solute Carrier Family 7, Member 11 (SLC7A11), solute carrier family 39, member 14 (SLC39A14), acyl-CoA synthetase long-chain family member 4 (ACSL4), cyclooxygenase-2(COX-2) ZO-1, Occludin, Claudin-1, Junctional adhesion molecule (JAM)1, JAM4 succinic semialdehyde dehydrogenase (SSADH), GABA transaminase (GABA-T), Gad1 and Gad2 and β-actin (primer sequences are listed in sFig. 5). Gene expression was determined using the SYBR Green kit (Vazyme, Q711-02). All the results were normalised against β-actin expression using the Thermal Cycler Dice Real Time System (ABI QuantStudio6, Singapore).

Revised version: Total RNA was extracted from liver and intestinal tissues using TRIzol (Yamei, YY101). Isolated RNA was reverse transcribed to cDNA using a kit (Vazyme, R323-01). The obtained cDNA was subjected to PCR using primers designed to detect Interleukin (Il)6, Il1b, Il18, recombinant solute carrier family 7, Member 11 (Slc7a11), solute carrier family 39, member 14 (Slc39a14), acyl-CoA synthetase long-chain family member 4 (Acsl4), cyclooxygenase- 2(Ptgs2), Tjp1, Ocln, Cldn7, Junctional adhesion molecule 1 (F11r), Junctional adhesion molecule 4 (Igsf5), succinic semialdehyde dehydrogenase (Aldh5a1), GABA transaminase (Abat), Gad1, Gad2 and Actb (primer sequences are listed in sFig. 5). Gene expression was determined using the SYBR Green kit (Vazyme, Q711-02). All the results were normalised against β-actin expression using the Thermal Cycler Dice Real Time System (ABI QuantStudio6, Singapore).

Western blotting

Total protein samples were extracted from tissues using RIPA lysis buffer (Yamei, PC101). Protein concentrations were determined using a BCA protein detection kit (Yamei, ZJ101). Proteins were separated using 10% SDS-PAGE and transferred to polyvinylidene fluoride membranes. Membranes were blocked in 5% skimmed milk then incubated with the primary antibodies ACSL4 (AbCam, ab155282), Transferrin Receptor 1 (TFR1) (Abclonl, A5865), ferritin heavy chain (FTH1) (AbCam, ab65080), Occludin (Proteintech, 27260-1-AP), GAPDH (Proteintech, 0494-1-AP) and β-tubulin (Proteintech, 10068-1-AP) at 4 °C overnight. After washing three times with tris buffered saline, the membranes were incubated with appropriate anti-rabbit or anti-mouse secondary antibodies at room temperature for 1 h. Imprinting was observed using chemiluminescence (Yamei, SQ201) and an Odyssey imaging system (Li-Cor- Biosciences, NE, USA).

High-performance liquid chromatography with tandem mass spectrometric detection (HPLC– MS/MS) detection of metformin and liver amino acids

Sample separation was performed using a ZORBAX Eclipse XDB-C18 column (4.6 × 150 mm, Agilent, USA) using an injection volume of 5 μL. Metformin detection conditions are as follows: column temperature of 30°C, one mobile phase is NaH2PO4 (Sigma, V900060) containing 5 mmol/L sodium dodecyl sulfonate (Sigma, V900859) with pH 3.5 while the other mobile phase is acetonitrile (Aladdin, A104440). The metformin content in each sample was measured by HPLC (Agilent 1260, USA). Amino acid detection conditions are as follows: column temperature of 40°C, mobile phase A of 10% methanol/water (containing 0.1% formic acid), mobile phase B of 50% methanol/water (containing 0.1% formic acid) and flow rate of 0.4 mL/min. Mass spectrometry was performed using a triple quadrupole mass spectrometer with an ESI source in negative ionisation mode. High-purity nitrogen was used as the nebulising and drying gas. Quantification was performed in multiple reaction monitoring mode.

16S rRNA sequencing

16S rRNA amplicon sequencing was performed at Chunlab Inc. (Seoul, Korea) with MiSeq system (Illumina). Briefly, for preparation of MiSeq library amplicons, target gene (16S rRNA V3–V4 region) was amplified using338F (5’-ACTCCTACGGGAGGCAGCAG-3’) and 806R (5’-GGACTACHVGGGTWTCTAAT-3’) primers, using an ABI GeneAmp 9700 PCR thermocycler (ABI, CA, USA). The PCR product was extracted from 2% agarose gel and purified using the AxyPrep DNA Gel Extraction Kit (Axygen Biosciences, Union City, CA, USA) according to the manufacturer’s instructions and quantified using a Quantus Fluorometer (Promega, USA). Normalization was performed the counts of individual OTUs in a sample by dividing the total counts of all OTUs within that sample followed by a multiplication by resulting in relative abundance expressed.

Metagenomic sequencing

Total genomic DNA was extracted from human fecal samples as mentioned in 16S rRNA sequencing. DNA extract was fragmented to an average size of about 400 bp using Covaris M220 (Gene Company Limited, China) for paired-end library construction. Paired-end library was constructed using NEXTflexTM Rapid DNA-Seq (Bioo Scientific, Austin, TX, USA). Adapters containing the full complement of sequencing primer hybridization sites were ligated to the blunt-end of fragments. Paired-end sequencing was performed on Illumina NovaSeq/Hiseq Xten (Illumina Inc., San Diego, CA, USA) at Majorbio Bio-Pharm Technology Co., Ltd. (Shanghai, China) using NovaSeq Reagent Kits/HiSeq X Reagent Kits according to the manufacturer’s instructions (www.illumina.com).

Statistical analysis

GraphPad Prism version 9.0 (GraphPad Software, San Diego, CA) was used for statistical treatment. Experimental data were shown as the mean ± SD. Two-tailed unpaired Student’s t- test and one-way ANOVA with Tukey’s correction were used for all comparisons of mice-related experiments. A P value < 0.05 was considered significant. The sample distribution was determined using a Kolmogorov–Smirnov normality test.